Subregional activity in the dentate gyrus is amplified during elevated cognitive demands

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

This manuscript investigates how dentate gyrus (DG) granule cell subregions, specifically suprapyramidal (SB) and infrapyramidal (IB) blades, are differentially recruited during a high cognitive demand pattern separation task. The authors combine TRAP2 activity labeling, touchscreen-based TUNL behavior, and chemogenetic inhibition of adult-born dentate granule cells (abDGCs) or mature granule cells (mGCs) to dissect circuit contributions.

This manuscript presents an interesting and well-designed investigation into DG activity patterns under varying cognitive demands and the role of abDGCs in shaping mGC activity. The integration of TRAP2-based activity labeling, chemogenetic manipulation, and behavioral assays provides valuable insight into DG subregional organization and functional recruitment. However, several methodological and quantitative issues limit the interpretability of the findings. Addressing the concerns below will greatly strengthen the rigor and clarity of the study.

Major points:

(1) Quantification methods for TRAP+ cells are not applied consistently across panels in Figure 1, making interpretation difficult. Specifically, Figure 1F reports TRAP+ mGCs as density, whereas Figure 1G reports TRAP+ abDGCs as a percentage, hindering direct comparison. Additionally, Figure 1H presents reactivation analysis only for mGCs; a parallel analysis for abDGCs is needed for comparison across cell types.

In Figure 1G and 1H we report TRAP+ abDGCs as a percentage rather than density because we are analyzing colocalization of the two markers, which are very sparse in this population. Given the very low number of double-labeled abDGCs, calculating density would not be practical. In the revised manuscript we have clarified the rationale for using these measures. As noted in the current text, we did not observe abDGCs co-expressing TRAP and c-Fos; we have made this point more explicit to guide interpretation of these data.

(2) The anatomical distribution of TRAP+ cells is different between low- and high-cognitive demand conditions (Figure 2). Are these sections from dorsal or ventral DG? Is this specific to dorsal DG, as itis preferentially involved in cognitive function? What happens in ventral DG?

The sections shown in Figure 2 were obtained from the dorsal dentate gyrus (see Methods, “Histology and imaging”: stereotaxic coordinates −1.20 to −2.30 mm relative to bregma, Paxinos atlas). From a feasibility standpoint, it is not possible to analyze the entire longitudinal extent of the hippocampus with these low-throughput histological approaches. We therefore focused on the dorsal DG, for which there is a strong functional rationale. A large body of work indicates that the dorsal hippocampus, and specifically the dorsal DG, is preferentially involved in spatial memory and in the fine contextual discrimination that underlies pattern separation. The dorsal hippocampus is critical for encoding and distinguishing similar spatial representations, a core component of the high-cognitive demand task used here. In contrast, the ventral DG is more strongly associated with emotional regulation and affective memory processing and is less implicated in high-resolution spatial encoding. For these reasons, the present study was designed to assess TRAP+ cell distributions specifically in the dorsal DG.

(3) The activity manipulation using chemogenetic inhibition of abDGCs in AsclCreER; hM4 mice was performed; however, because tamoxifen chow was administered for 4 or 7 weeks, the labeled abDGC population was not properly birth-dated. Instead, it consisted of a heterogeneous cohort of cells ranging from 0 to 5-7 weeks old. Thus, caution should be taken when interpreting these results, and the limitations of this approach should be acknowledged.

We agree that prolonged tamoxifen administration results in labeling a heterogeneous population of abDGCs spanning approximately 0 to 5–7 weeks of age, rather than a precisely birth-dated cohort. This is a limitation of this approach and we have included discussion of this in more detail in the revised manuscript.

(4) There is a major issue related to the quantification of the DREADD experiments in Figure 4, Figure 5, Figure 6, and Figure 7. The hM4 mouse line used in this study should be quantified using HA, rather than mCitrine, to reliably identify cells derived from the Ascl lineage. mCitrine expression in this mouse line is not specific to adult-born neurons (off-targets), and its expression does not accurately reflect hM4 expression.

We agree that mCitrine is not a marker that allows localization of hM4Di as it is well known that the mCitrine can be independently expressed in a Cre independent manner in this mouse. As suggested, we have removed the figure that showed the mCitrine and have performed immunohistochemical localization of the DREADD with an antibody against the HA tag. This is now shown in Figure 5.

(5) Key markers needed to assess the maturation state of abDGCs are missing from the quantification. Incorporating DCX and NeuN into the analysis would provide essential information about the developmental stage of these cells.

The goal of this study was to examine activity patterns of adult-born versus mature granule cells, rather than to assess maturation state. The adult-born neurons analyzed were 25–39 days old, an age at which point most cells have progressed beyond the DCX+ stage and are expected to express NeuN based on prior work. We therefore do not think that including DCX or NeuN quantification would provide additional information relevant to the aims or interpretation of this study.

Minor points:

(1) The labeling (Distance from the hilus) in Figure 2B is misleading. Is that the same location as the subgranular zone (SGZ)? If so, it's better to use the term SGZ to avoid confusion.

We have updated Figure 2B, the Methods, and the main text to more explicitly localize this which it the boundary between the subgranular zone (SGZ) and the hilus.

(2) Cell number information is missing from Figures 2B and 2C; please include this data.

We have now added the cell number information to the figure legends. In Figures 2B and 2C, each point corresponds to a single cell, with an equal number of mice per group. The total number of TRAP+ cells per mouse is shown in Figure 1F, which reports TRAP+ cell densities by group.

(3) Sample DG images should clearly delineate the borders between the dentate gyrus and the hilus. In several images, this boundary is difficult to discern.

We made the DG-hilus boundaries clearer in the sample images to improve visualization and interpretation.

(4) In Figure 6, it is not clear how tamoxifen was administered to selectively inhibit the more mature 6-7-week-old abDGC population, nor how this paradigm differs from the chow-based approach. Please clarify the tamoxifen administration protocol and the rationale for its specificity.

We apologize for the confusion here. The protocol used in Figure 6 is the same tamoxifen chow–based approach as in Figure 5, differing only in the duration of tamoxifen exposure. Mice in Figure 5 received tamoxifen chow for 7 weeks, whereas mice in Figure 6 received it for 4 weeks, restricting labeling to a younger and narrower cohort of adult-born DGCs. Thus, the population targeted in Figure 6 is younger than that in Figure 5 and does not correspond to mature 6–7-week-old neurons. By contrast, the experiment in Figure 4 targets a more mature population, consisting predominantly of ~5-week-old adult-born neurons as well as mature granule cells, which are Dock10-positive and express Cre endogenously, allowing selective manipulation of this later-stage population.

We have corrected the paragraph accordingly and clarified the age range of the labeled populations in the revised manuscript.

Comments on revisions:

I appreciate the authors' careful and thorough revisions. They have addressed all of my previous concerns satisfactorily, and the manuscript is now significantly strengthened. I have no further concerns.

Reviewer #2 (Public review):

In this study, the authors investigate how increasing cognitive demand shapes activity patterns in the dorsal dentate gyrus (DG). Using a touchscreen-based TUNL task combined with TRAP/c-Fos tagging, birth-dating of adult-born granule cells (abDGCs), and chemogenetic inhibition, they show that higher task demand increases mature granule cell (mGC) recruitment and enhances suprapyramidal (SB) versus infrapyramidal (IB) blade bias. Functionally, mGC inhibition reduces overall activity and impairs performance without disrupting blade bias, whereas inhibition of {less than or equal to}7-week-old abDGCs increases mGC activity, abolishes blade bias, and impairs discrimination under high-demand conditions. These findings suggest that effective pattern separation depends not only on overall DG activity levels but also on the spatial organization of recruited ensembles.

The integration of touchscreen TUNL with temporally controlled activity tagging and birth-dated cohorts is technically strong. Quantification of SB-IB bias and radial/apical distributions adds anatomical precision beyond bulk activity measures. The comparison between mGC and abDGC inhibition is conceptually compelling and supports dissociable functional roles. Overall, the data convincingly demonstrate that increasing cognitive demand amplifies blade-biased DG recruitment and that mGCs and abDGCs differentially contribute to both behavioral performance and network organization.

However, how abDGCs are integrated into the mGC network under high cognitive demand remains unresolved. Additional experiments are needed to clarify how abDGCs shape spatial recruitment patterns and whether they directly inhibit or indirectly regulate mGC activity to maintain high performance.

Furthermore, the authors frame "high cognitive demand" as a multidimensional construct encompassing broad behavioral challenge. It would strengthen the work to delineate how local abDGC-mGC circuit interactions regulate specific task components in real time. This will require higher temporal resolution approaches, as TRAP and c-Fos labeling integrate activity over prolonged windows and primarily reflect sustained engagement rather than moment-to-moment computations.

The central conclusion that dentate function depends on coordinated spatial recruitment rather than total activity magnitude is supported by the data, although mechanistic interpretations should be tempered given methodological limitations.

Overall, this work advances models of adult neurogenesis by emphasizing a critical-period modulatory role of abDGCs in organizing DG network activity during high-demand discrimination. The combined behavioral and circuit-level framework is likely to be influential in the field.

Reviewer #3 (Public review):

This study examines the role of dentate gyrus neuronal populations, reflecting neurogenesis and anatomical location (suprapyramidal vs infrapyramidal blade), in a mnemonic discrimination task that taxes the pattern separation functions of the dentate. The authors measure dentate gyrus activity resulting from cognitive training and test whether adult neurogenesis is required for both the anatomical patterns of activity and performance in the cognitive task. The authors find that more cognitively challenging variants of the task evoked more dentate activity, but also distinct patterns of activity (more activity in the suprapyramidal blade, less in the infdrapyramidal blade). Using chemogenetic approaches they silence mature vs immature dentate gyrus neurons and find that only mature neurons (either the general population or specifically mature adult-born neurons), and not immature adult-born neurons, are required for the difficult version of the task. Inhibition of mature adult-born neurons furthermore increased overall activity in the dentate and reduced the biased pattern of activity across the blades, consistent with evidence that adult-born neurons broadly regulate dentate gyrus activity.

Comments on revisions:

I appreciate the efforts the authors have taken to revise this manuscript. I have only minor concerns with this revised version of the manuscript:

Methods state that significance is defined as P<0.05 but some results are interpreted as significant when P=0.05. Either the alpha value needs to change or the interpretation needs to change.

We have corrected the statement in the Methods section to define statistical significance as P ≤ 0.05, which aligns with how significance was interpreted throughout the manuscript.

I believe the statistical results for group and blade effects for the ANOVAs, in Figs 2,3 & 4, appear to be switched (blade should be significant, not group).

We thank the reviewer for pointing out this mistake. We have corrected the reported statistical results for the group and blade effects in the manuscript accordingly.

I appreciate that sometimes there is not a perfect overlap between immunohistochemical signals, but I continue to believe that the spatially-non-overlapping TRAP and EDU signals in Fig 3 is caused by these 2 markers being in different cells. A Z-stack or orthogonal projection could verify/disprove this concern.

We agree that limited overlap in single optical sections can raise the possibility that TRAP and EdU signals originate from different cells. However, based on our imaging conditions and inspection across focal planes, the signals are consistent with being present within the same cells, with partial spatial separation likely reflecting subcellular localization and/or sectioning effects.

Changes from the earlier version to the revised version

Abstract

Earlier version: Neural activity in the dentate gyrus (DG) is required for the detection and discrimination of novelty, context and patterns, amongst other cognitive processes. Previous studies have shown that granule cell activation differs between the supra and infrapyramidal blades of the DG during a range of hippocampal dependent tasks yet how excitatory dynamics within the DG drive and modulate this blade-specific bias under varying cognitive demands remains unclear. Here we used an automated touch screen pattern separation task combined to temporally controlled tagging of active neurons to determine how increasing cognitive demand shapes the spatial patterns of activity in the DG. We observed that as task difficulty increased, DG activation became progressively biased toward the suprapyramidal blade, accompanied by structured distributions of active mature granule cells (mGCs) along the apex-to-blade and hilar-to-molecular layer axes. Selective inhibition of mature granule cells (mGCs) did not affect the patterned distribution of active cells but profoundly impaired performance, as mice were no longer able to discriminate between closely spaced locations. In contrast, chemogenetic inhibition of adult-born granule cells (abDGCs) beyond a critical window of their maturation significantly impaired performance of mice during high-demand conditions, elevated overall mGC activity and disrupted the blade-specific distribution of active mGCs even in animals that successfully completed the task. These findings demonstrate how a high cognitive demand pattern separation task preferentially activates mGCs in subregions of the DG and are consistent with a modulatory role for abDGCs on the dentate circuit which in part governs the spatially organized patterns of activity of mGCs.

Revised version: Neural activity in the dentate gyrus (DG) supports the detection and discrimination of novelty, context, and patterns. Granule cell activation differs between the supra-and infrapyramidal blades across hippocampal-dependent tasks, yet how excitatory dynamics shape this blade-specific bias under varying cognitive demands remains unclear. Here, we combined an automated touchscreen pattern separation task in mice with temporally controlled tagging of active neurons to determine how increasing cognitive demand influences spatial activity patterns in the DG. As task difficulty increased, activation became progressively biased toward the suprapyramidal blade and was accompanied by structured distributions of active mature granule cells (mGCs) along both the radial and transverse axes. Selective inhibition of mGCs did not alter these spatial patterns but profoundly impaired performance, as mice were no longer able to discriminate between closely spaced locations. In contrast, chemogenetic inhibition of adult-born granule cells (abDGCs) beyond a critical maturation window impaired performance under high-demand conditions, increased overall mGC activity, and disrupted blade-specific organization even in animals that successfully completed the task. These findings demonstrate that high cognitive demand recruits spatially organized mGC activity and support a modulatory role for abDGCs in shaping dentate circuit dynamics.

Introduction

The dentate gyrus (DG) is the input structure to the hippocampus and plays important roles in episodic memory1, spatial coding2, novelty detection3 as well as in separating patterns4. A long-standing proposal stemming from early computational studies, suggest that the DG contribute to pattern separation based on its anatomical connectivity5–7. The DG is the gateway to the hippocampus receiving direct excitatory input from the entorhinal cortex (EC). The large number of granule neurons in the DG receiving input from a relatively smaller population of EC has been proposed to promote expansion recoding which, along with sparse firing of granule cells, allows the separation of inputs at the level of neural activity8.

The granule cell layer of the DG is organized into two blades, termed the suprapyramidal (SB) and infrapyramidal (IB) blades. When viewed in transverse cross-section9, mature granule cells (mGCs) exhibit asymmetric activation across these blades during hippocampal-dependent tasks10–13, with greater activity typically observed in the SB compared to the IB10–12 However, the circuit-level mechanisms underlying this functional asymmetry remain unresolved.

In addition to mature populations, the subgranular zone (SGZ) of the DG serves as a neurogenic niche where new granule neurons continue to be generated throughout adulthood14–17. This process of adult neurogenesis gives rise to adult-born dentate granule cells (abDGCs), which can remodel local circuitry and potentially contribute to memory flexibility18. Interestingly, the asymmetry in mGC activation across the blades parallels an asymmetry in neurogenesis, with higher numbers of abDGCs reported in the SB19–22.

In this study, we investigated whether performance in a cognitively demanding task alters the patterns of neural activity in the dorsal dentate gyrus (DG), and how manipulation of excitatory populations within the DG can modulate these activity patterns. Using TRAP2 mice23 to label neurons activated during a touchscreen-based spatial discrimination task, we observed a consistent bias in activity between the two blades of the DG. When the cognitive demand of the task was elevated by reducing the degree of spatial separation, there was an increase of TRAP cells activity and a more prominent and amplified spatial bias with active neurons distributed preferentially to the SB and relatively reduced in the IB. To dissect the contribution of distinct excitatory populations, we selectively inhibited either mature or adult-born granule cells during task performance. We found that manipulating these populations had differential effects on the blade-specific bias of activation, indicating distinct roles for mGCs and abDGCs in shaping DG representations during spatial cognition. Inhibition of mature granule cells did not alter the blade bias of activation but markedly impaired performance, as mice were no longer able to discriminate closely spaced locations under high cognitive demand. In contrast, chemogenetic inhibition of abDGCs beyond a critical window of their maturation, decreased task performance, increased overall mGC activity, and reduced the blade-specific bias of activation, suggesting that abDGC activity is required to maintain proper DG network organization under high cognitive load. Together, these findings reinforce the significance of neural activity in the dorsal DG in spatial pattern separation and highlight the circuit contributions of both abDGCs and mGCs to behaviors requiring a high cognitive demand.

Results

Amplified blade-biased activity of GCs during a high cognitive demand pattern separation task

Earlier version: Prior work has demonstrated that expression of immediate early genes in the DG including Arc10, c-Fos12 and Zif26824 correlate with behavioral state and are active in both mature and young abDGCs. Genetic strategies using TRAP1 mice25 and TetTag mice26,27 have also demonstrated that active GC populations can be labeled with these strategies. To determine whether behaviorally induced labeling of both mGCs and abDGCs is observed in TRAP2 mice (Fos2a-iCreERT2;Ai9) we combined birth-dating of abDGCs with a single exposure to voluntary wheel running which is known to activate DG neurons25 (Figure S1A). Mice given access to running wheels for 4 hours had elevated numbers of TRAP+ neurons throughout the DG, with both mature (mGCs; defined as granule cells not birth-dated with BrdU or EdU) and immature neurons labeled (Figure S1B-F). TRAP+ labeling of abDGCs was observed in birth-dated neurons 14 days post injection (dpi) from BrdU (Figure S1E & F). No double labeled TRAP+ BrdU+ neurons were observed in younger neurons prior to this timepoint (Figure S1F).

Revised version: Prior work has demonstrated that expression of immediate early genes in the DG including Arc 10, c-Fos12 and Zif26824 correlate with behavioral state and are active in both mature and young abDGCs. Genetic strategies using TRAP1 mice25 and TetTag mice26,27 have also demonstrated that active GC populations can be labeled with these strategies. To determine whether behaviorally induced labeling of both mGCs and abDGCs is observed in TRAP2 mice (Fos2a-iCreERT2;Ai9) we combined birth-dating of abDGCs with a single exposure to voluntary wheel running which is known to activate DG neurons25 (Figure S1A). Mice given access to running wheels for 4 hours had elevated numbers of TRAP+ neurons throughout the DG, with both mature (mGCs; defined as granule cells not birth-dated with BrdU or EdU) and immature neurons labeled (Figure S1B-F). TRAP+ labeling of abDGCs was observed in birth-dated neurons 14 days post injection (dpi) from BrdU (Figure S1E & F). No double labeled TRAP+ BrdU+ neurons were observed in younger neurons prior to this timepoint (Figure S1F).

Importantly, an analysis of the blade specific localization of TRAP+ neurons demonstrated that in home cage controls there was a small bias in the distribution of active neurons to the SB of the DG and this distribution of labeled GCs was not different in mice exposed to a single episode of running (Figure S1D). Moreover, consistent with previous reports19–21, we observed an asymmetry in neurogenesis, with a greater number of abDGCs in the SB compared to the IB (Figure S1G).

In order to determine the activation patterns of GCs (mGCs and abDGCs) when mice successfully performed a high cognitive demand spatial discrimination pattern separation task, TRAP2 mice were trained in a modified touch screen task that tests the ability to separate a unique pattern of illuminated squares posing a relatively low or high cognitive demand (Figure 1A and Figure S2A & B). The trial unique nonmatching-to-location (TUNL) task has been demonstrated to require the hippocampus, particularly when the cognitive demand is high when the separation of the sample to test squares is reduced, or the delay between sample and choice is increased28,29. Mice were injected with BrdU to birthdate neurons at day 0 (D0). At D7 mice were water restricted before EdU injections and shaping at D14 to habituate them to the operant task and learn to touch the screen for a saccharine water reward. Mice were then trained until they reached criterion (70% correct during one of two daily sessions on 2 consecutive days) in the TUNL task with a 1s delay between the sample and choice and a large (L) separation between sample panel and the test panel (Figure 1A and Figure S2A). Mice reached criterion on average at 10.7 ± 0.8 days in the large separation sessions (Figure 1B & D left panel). After reaching criterion the mice were tested with alternating trials of large and small (LS) separation between the sample and test panels. In this case the mice were always above criterion for the L trials but took another 9.3 ± 1.0 days to reach criterion on the S trials (Figure 1C & D middle and right panel).

TRAP labeling of mGCs and birth-dated abDGCs after the TUNL pattern separation task

(A) Top: Schematic illustrating the pattern separation paradigm. Male (n = 9) and female (n = 7) mice (6-8-week-old) received 3 injections of BrdU at D0 and 3 injections of EdU two weeks later (D14) before starting the behavioral paradigm. Mice received one injection of 4-OHT on the day they reached criterion (TRAP) and were perfused 3 days later 90 minutes after one exposure to the behavioral paradigm to allow endogenous c-Fos expression (c-Fos). Control mice (males n = 6, females n = 4) underwent the same treatment but only participated in the Large separation task (L). Mice were water-restricted during the behavioral experiment. Bottom: Cartoon representation of the chamber in TUNL task with examples of Large separation (L) and Small (S) separation configurations. High cognitive demand paradigm consisted of interleaved L and S trials (LS). (B) Number of days for each mouse to reach the 70% success criterion in the Large separation (L). (C) Number of days for mice to reach 70% success criterion in S trials, during the LS separation task. (D) Left: Average percentage of correct choices (% success) in the L separation on the first day of training, as well as during the three days prior to reaching criterion (days before criterion, DBC) and the day of reaching the criterion (crit)(one-sample t-test to chance at criterion Large t(16)=12.04, p<0.0001). Middle: Average percentage of correct choices in the L and S separation on the first day of training, during the three days before criterion (DBC), on the day of reaching criterion and 3 days later (one-sample t-test to chance at criterion Large t(16)=8.38, p<0.0001; Small t(16)=7.37, p<0.0001). Right: Cumulative probability of days to criterion of mice in S for both controls L and LS groups. (E) Example section of BrdU+ labeling (green) EdU+ labeling (blue) and TRAP+ labeling (red) in TRAP2 mice. Calibration bar: 30 μm. (F) Density of TRAP-labeled GCs in mice undergoing LS compared to control mice engaged in either shaping (Shap) or the Large configuration task only (L) (one-way ANOVA, F(2,25) = 7.97, p = 0.0021; Tukey post-hoc test. Groups Shap vs LS and Large vs LS were significantly different). (G) Percentage of BrdU labeled cells (39 days +/-1.23) or Edu labeled cells (25 days +/-1.37 days) also TRAP+ after L or LS training (BrdU+/TRAP+ L: 0.145 ± 0.145 %, 6 mice; LS BrdU+/TRAP+ 0.822 ± 0.474 %, 6 mice; ns; EdU+/TRAP+ L: 0.203 ± 0.144 %, 6 mice; LS BrdU+/TRAP+ 0.304 ± 0.155 %, 7 mice; ns; Mann–Whitney). (H) Percentage of TRAP+ GCs expressing c-Fos+ in Shap, L or LS task after reaching criterion (expressed as a percentage of TRAP+ cells) (one-way ANOVA, F(2,18) = 6.66, p = 0.0069; Tukey post-hoc test. Groups Shap vs LS and Large vs LS were significantly different).

Different cohorts of mice were analyzed at each level of cognitive demand of the task to assess the activation of mGCs and abDGCs. TRAP labeling was used to identify active neurons in the DG, together with thymidine analogs that birth-dated abDGCs at approximately 3 weeks (25 ± 1.4 days, EdU) and 5 weeks (39 ± 1.2 days, BrdU) of age (Figure 1E). We found a strong relationship between the level of cognitive demand of the task and the degree of mGC activation with a greater number of TRAP+ neurons in mice tested under higher cognitive load.

Specifically, TRAP+ neuron density increased from shaping to the L separation stage, and was highest in mice performing the most cognitively demanding condition (LS) (one-way ANOVA, F(2,25) = 7.97, p = 0.0021; Tukey post-hoc test. Groups Shap vs LS and Large vs LS were significantly different) (Figure 1F), regardless of whether the mice reach the criterion (C) or not (NC)(Figure S2H). A similar trend was observed when assessing c-Fos labeling 90 minutes after the final training session, further supporting the link between task demand and mGC activation (Figure S2F). Furthermore, when we examined the correlation between successful performance to criterion and the number of TRAP+ neurons, we found no significant relationship (Figure S2C), demonstrating that increased DG activity is driven by engagement in a high-demand task rather than task success.

Earlier version: Activation of adult-born granule cells (abDGCs) birth-dated with BrdU and EdU was assessed by colocalization with TRAP+ cells, rather than density, because the low number of newborn neurons precluded reliable density measurements. Accordingly, results are expressed as the proportion of birth-dated cells that were TRAP+. Using this approach, we did not observe a relationship between the level of cognitive demand and the recruitment of newborn neurons. Only a small fraction (∼1%) of 25- to 39-day-old abDGCs were active (TRAP+) during the high cognitive demand task (LS) (Figure 1G), regardless of whether the mice reached criterion (C) or not (NC) (Figure S2J). Given the sparse labeling in this population, this percentage reflects very low absolute cell numbers per animal. In contrast, mice that performed only the low-demand version of the task (L configuration) exhibited almost no TRAP-labeled birth-dated cells (Figure 1G). Moreover, in the shaping group, no newborn neurons were co-labeled with TRAP, indicating that minimal cognitive engagement does not recruit this population (Figure 1G). Importantly, we did not detect qualitative differences in the overall abundance or blade distribution of BrdU+ or EdU+ neurons across behavioral groups (not shown).

Revised version: Activation of adult-born granule cells (abDGCs) birth-dated with BrdU and EdU was assessed by colocalization with TRAP+ cells, rather than density, because the low number of newborn neurons precluded reliable density measurements. Accordingly, results are expressed as the proportion of birth-dated cells that were TRAP+. Using this approach, we did not observe a relationship between the level of cognitive demand and the recruitment of newborn neurons. Only a small fraction (∼1%) of 25-to 39-day-old abDGCs were active (TRAP+) during the high cognitive demand task (LS) (Figure 1G), regardless of whether the mice reached criterion (C) or not (NC) (Figure S2J). Given the sparse labeling in this population, this percentage reflects very low absolute cell numbers per animal. In contrast, mice that performed only the low-demand version of the task (L configuration) exhibited almost no TRAP-labeled birth-dated cells (Figure 1G). Moreover, in the shaping group, no newborn neurons were co-labeled with TRAP, indicating that minimal cognitive engagement does not recruit this population (Figure 1G).

Added in the revised version: Importantly, we did not detect qualitative differences in the overall abundance or blade distribution of BrdU+ or EdU+ neurons across behavioral groups (not shown).

In a separate cohort, we TRAPed neurons during the final day of LS trials and re-exposed mice to the same task 3 days later. After two LS sessions, mice were perfused 90 minutes later for immunohistochemical detection of endogenous c-Fos (Figure 1A). Task performance was reduced compared to the final day of prior training in both the L and S trials (Figure 1D, middle panel). Analysis of mGC reactivation (TRAP+/c-Fos+ cells) revealed greater overlap in mice performing LS trials than in those performing only L separations or shaping sessions (one-way ANOVA, F(2,18) = 6.66, p = 0.0069; Tukey post-hoc test. Groups Shap vs LS and Large vs LS were significantly different) (Figure 1H). Notably, across all animals and conditions, no BrdU⁺ or EdU⁺ cells were found to be labelled with both TRAP and c-Fos, indicating that newborn neurons activated during the final LS session were not reactivated three days later. Collectively, these results indicate that while increasing cognitive demand enhances activation and reactivation of mature granule cells, the recruitment of adult-born neurons remains low and does not scale with task difficulty.

Earlier version: Consistent with previous reports10–13, we observed a bias in granule cell activity within the DG, with TRAP+ mGCs preferentially localized to the suprapyramidal blade (SB). Interestingly, this bias was not fixed but varied according to the level of cognitive demand of the task (Figure 2A). Mice performing the high-demand LS condition showed a pronounced enrichment of TRAP+ mGCs in the SB compared to the infrapyramidal blade (IB), driven by both an increase in labeled cells in the SB and a decrease in the IB (two-way ANOVA (unbalanced) (group: F(1,38) = 146.34, p = 1.33 × 10⁻¹⁴) (blade: F(2,38) = 1.14 × 10⁻⁷, p = 1) (group × blade: F(2,38) = 13.99, p = 2.81 × 10⁻⁵, Tukey post-hoc test, LS and Large (p=0.03), LS and Shap (p=0.05) for SB and IB) (Figures 2A and S2E). In contrast, mice performing lower-demand versions of the task still exhibited a significant SB–IB difference, but the magnitude of this bias was reduced, resulting in a more balanced distribution of activity across blades. However, as with overall TRAP+ cell density, there was no significant correlation between individual task performance and the distribution of TRAP+ neurons in the SB (Figure S2D).

Revised version: Consistent with previous reports10–13, we observed a bias in granule cell activity within the DG, with TRAP+ mGCs preferentially localized to the suprapyramidal blade (SB). Interestingly, this bias was not fixed but varied according to the level of cognitive demand of the task (Figure 2A). Mice performing the high-demand LS condition showed a pronounced enrichment of TRAP+ mGCs in the SB compared to the infrapyramidal blade (IB), driven by both an increase in labeled cells in the SB and a decrease in the IB (two-way ANOVA (unbalanced) (blade: F(1,38) = 146.34, p = 1.33 × 10⁻¹⁴) (group: F(2,38) = 1.14 × 10⁻⁷, p = 1) (group × blade: F(2,38) = 13.99, p = 2.81 × 10⁻⁵, Tukey post-hoc test, LS and Large (p=0.03), LS and Shap (p=0.05) for SB and IB) (Figures 2A and S2E). In contrast, mice performing lower-demand versions of the task still exhibited a significant SB–IB difference, but the magnitude of this bias was reduced, resulting in a more balanced distribution of activity across blades. However, as with overall TRAP+ cell density, there was no significant correlation between individual task performance and the distribution of TRAP+ neurons in the SB (Figure S2D).

Blade-biased activity of mGCs during a high cognitive demand pattern separation task.

Earlier version: (A) Spatial distribution of TRAP-labeled cells in the infrapyramidal (IB) and suprapyramidal (SB) blades of the dorsal DG in TRAP2 mice performing the LS separation, compared to control mice that performed only the Shap or L configuration (two-way ANOVA (unbalanced) (group: F(1,38) = 146.34, p = 1.33 × 10⁻¹⁴) (blade: F(2,38) = 1.14 × 10⁻⁷, p = 1) (group × blade: F(2,38) = 13.99, p = 2.81 × 10⁻⁵). Tukey post-hoc test: significant differences were observed between IB; LS and Large (p=0.03), LS and Shap (p=0.05), SB; LS and Large (p= 0.03), LS and Shap (p=0.05). (B) Spatial distribution of TRAP+ cells in the hilar to outer radial GCL axes in the SB (0 to 120μm) and the IB (0 to 120μm) of the dorsal DG with 0 μm indicating location at the hilar border (Shap: 5 mice, 255 cells; L: 5 mice, 294 cells; LS: 5 mice, 639 cells). (C) Distribution of TRAP-labeled cells along the apex to blade extremity axes of the SB of the dorsal DG, with 0% indicating location at the dentate apex (Shap: 5 mice, 152 cells; L: 5 mice, 176 cells; LS: 5 mice, 517 cells)( two-way ANOVA (unbalanced) (group: F(1,26) = 26.05, p = 2.56 × 10⁻⁵) (distance (0%-100%): F(2,26) = 0.044, p = 0.957) (group× distance (0%-100%): F(2,26) = 23.77, p = 1.35 × 10⁻⁶). Tukey post-hoc test: significant differences were observed between Shap; 0%-50% and 50%-100% (p<0.0001), LS; 0%-50% and 50%-100% (p=0.03). (D) Example section of TRAP+ cells in the dorsal DG illustrating the distinct boundaries between the two blades, the SGZ, and the apex. Calibration bar: 50 μm. (E) Cartoon illustrating the two blades of the DG: IB and SB. Dentate granule cells are depicted in grey while red circles represent TRAP+ cells in the DG. In the L group, a greater number of labeled cells are localized in the SB compared to the IB. This bias of activity is also observed in the LS group, where the bias is more pronounced, and the overall activity of the DG is increased. In the SB, the majority of the activated cells are located closer to the tip of the blade.

Revised version: (A) Spatial distribution of TRAP-labeled cells in the infrapyramidal (IB) and suprapyramidal (SB) blades of the dorsal DG in TRAP2 mice performing the LS separation, compared to control mice that performed only the Shap or L configuration (two-way ANOVA (unbalanced) (blade: F(1,38) = 146.34, p = 1.33 × 10⁻¹⁴) (group: F(2,38) = 1.14 × 10⁻⁷, p = 1) (group × blade: F(2,38) = 13.99, p = 2.81 × 10⁻⁵). Tukey post-hoc test: significant differences were observed between IB; LS and Large (p=0.03), LS and Shap (p=0.05), SB; LS and Large (p= 0.03), LS and Shap (p=0.05). (B) Spatial distribution of TRAP+ cells in the hilar to outer radial GCL axes in the SB (0 to 120μm) and the IB (0 to 120μm) of the dorsal DG with 0 μm indicating location at the hilar border (Shap: 5 mice, 255 cells; L: 5 mice, 294 cells; LS: 5 mice, 639 cells). (C) Distribution of TRAP-labeled cells along the apex to blade extremity axes of the SB of the dorsal DG, with 0% indicating location at the dentate apex (Shap: 5 mice, 152 cells; L: 5 mice, 176 cells; LS: 5 mice, 517 cells)(two-way ANOVA (unbalanced) (distance (0%-100%): F(1,26) = 26.05, p = 2.56 × 10⁻⁵), group: F(2,26) = 0.044, p = 0.957) (group× distance (0%-100%): F(2,26) = 23.77, p = 1.35 × 10⁻⁶). Tukey post-hoc test: significant differences were observed between Shap; 0%-50% and 50%-100% (p<0.0001), LS; 0%-50% and 50%-100% (p=0.03). (D) Example section of TRAP+ cells in the dorsal DG illustrating the distinct boundaries between the two blades, the SGZ, and the apex. Calibration bar: 50 μm. (E) Cartoon illustrating the two blades of the DG: IB and SB. Dentate granule cells are depicted in grey while red circles represent TRAP+ cells in the DG. In the L group, a greater number of labeled cells are localized in the SB compared to the IB. This bias of activity is also observed in the LS group, where the bias is more pronounced, and the overall activity of the DG is increased. In the SB, the majority of the activated cells are located closer to the tip of the blade.

Earlier version: Further spatial mapping of TRAP+ cells relative to the SGZ and dentate apex revealed that active neurons were concentrated in the outer radial portion of the granule cell layer—an area enriched in spatially tuned semilunar granule cells 11,30, and this localization was particularly evident after the high-demand LS task (Figures 2B and 2C). In addition, active neurons were preferentially located toward the distal tip of the suprapyramidal blade (two-way ANOVA (unbalanced) (group: F(1,26) = 26.05, p = 2.56 × 10⁻⁵) (distance (0%-100%): F(2,26) = 0.044, p = 0.957) (group× distance (0%-100%): F(2,26) = 23.77, p = 1.35 × 10⁻⁶). Tukey post-hoc test: significant differences were observed between Shap; 0%-50% and 50%-100% (p<0.0001), LS; 0%-50% and 50%-100% (p=0.03)(Figure 2C).Together, these results demonstrate that the suprapyramidal bias in granule cell activity is dynamically modulated by cognitive demand, with the separation becoming more pronounced as task difficulty increases and producing distinct distributions of active neurons along both the apex–distal and hilar–outer axes of the DG (Figures 2D and 2E).

Revised version: Further spatial mapping of TRAP+ cells relative to the SGZ and dentate apex revealed that active neurons were concentrated in the outer radial portion of the granule cell layer—an area enriched in spatially tuned semilunar granule cells 11,30, and this localization was particularly evident after the high-demand LS task (Figures 2B and 2C). In addition, active neurons were preferentially located toward the distal tip of the suprapyramidal blade (two-way ANOVA (unbalanced) (distance (0%-100%): F(1,26) = 26.05, p = 2.56 × 10⁻⁵) (group: F(2,26) = 0.044, p = 0.957) (group× distance (0%-100%): F(2,26) = 23.77, p = 1.35 × 10⁻⁶). Tukey post-hoc test: significant differences were observed between Shap; 0%-50% and 50%-100% (p<0.0001), LS; 0%-50% and 50%-100% (p=0.03)(Figure 2C).Together, these results demonstrate that the suprapyramidal bias in granule cell activity is dynamically modulated by cognitive demand, with the separation becoming more pronounced as task difficulty increases and producing distinct distributions of active neurons along both the apex–distal and hilar–outer axes of the DG (Figures 2D and 2E).

Activation of GCs during remote recall in the high cognitive demand task

Earlier version: We next investigated whether a similar spatially biased activity of GCs was observed in a remote recall (R) test three weeks after reaching the criterion during the initial training (I), and whether there was overlap in the population of active DG neurons (Figure 3A). In the remote session, mice had a high success rate performing close to criterion in the L trials on the first day of retesting suggesting that 3 weeks after initial learning they retained the ability in task (Figure 3B). However, when confronted with the more demanding trials with inclusion of the small separation configuration, mice required additional days to achieve the criterion level (Figure 3C), although most mice required less time to reach criterion in LS training during the remote trials than on the initial testing (Figure 3D). We again observed that the density of TRAP+ cells was increased in mice that underwent LS compared to L only trials during the initial training (not shown) and importantly the distribution of active neurons during both the initial training (TRAP+) and remote training (c-Fos+) were biased to the SB after LS training (Figure 3F & G)(Two-way ANOVA: group, F(1,12) = 32.98, p < 0.001; blade, F(1,12) = 9.2×10⁻³⁰, p = 1; group×blade interaction, F(1,12) = 16.92, p = 0.0014. Tukey post-hoc test: significant differences were observed between IB; LS and Large (p=0.05), SB; LS and Large (p= 0.05)(3G: Two-way ANOVA: group, F(1,12) = 56.03, p < 0.00001; blade, F(1,12) = 0, p = 1; group×blade interaction, F(1,12) = 7.02, p = 0.021.Tukey post-hoc test: no significant differences were observed between LS and Large for IB or SB).

Revised version: We next investigated whether a similar spatially biased activity of GCs was observed in a remote recall (R) test three weeks after reaching the criterion during the initial training (I), and whether there was overlap in the population of active DG neurons (Figure 3A). In the remote session, mice had a high success rate performing close to criterion in the L trials on the first day of retesting suggesting that 3 weeks after initial learning they retained the ability in task (Figure 3B). However, when confronted with the more demanding trials with inclusion of the small separation configuration, mice required additional days to achieve the criterion level (Figure 3C), although most mice required less time to reach criterion in LS training during the remote trials than on the initial testing (Figure 3D). We again observed that the density of TRAP+ cells was increased in mice that underwent LS compared to L only trials during the initial training (not shown) and importantly the distribution of active neurons during both the initial training (TRAP+) and remote training (c-Fos+) were biased to the SB after LS training (Figure 3F & G)(Two-way ANOVA: blade: F(1,12) = 32.98, p < 0.001; group: F(1,12) = 9.2×10⁻³⁰, p = 1; group×blade interaction, F(1,12) = 16.92, p = 0.0014. Tukey post-hoc test: significant differences were observed between IB; LS and Large (p=0.05), SB; LS and Large (p= 0.05)(3G: Two-way ANOVA: blade: F(1,12) = 56.03, p < 0.00001; group: F(1,12) = 0, p = 1; group×blade interaction, F(1,12) = 7.02, p = 0.021.Tukey post-hoc test: no significant differences were observed between LS and Large for IB or SB).

Activation of GCs during remote recall in the high cognitive demand task.

Earlier version: (A) Schematic representation of the pattern separation paradigm used in the study. Mice (male n=11, female n=8) performed two trainings sessions: initial (I) and remote (R). The second training (R) took place 3 weeks after the mice reached the S criterion in LS. Mice were perfused when they reached the criterion in S during R. A second group that performed only in the L configuration during both training sessions was included as control (males n=3, females n=3). (B) Learning curves representing the percentage of success during the remote test (R). The control group trained only in the large-separation condition (L group) is shown in green and includes performance in the large-separation trials only. The LS group (orange) was trained in both large- and small-separation conditions and therefore includes two traces: the upper orange trace represents performance during large-separation trials, and the lower orange trace represents performance during small-separation trials. (one-sample t-test to chance at criterion; LS: Large t(14) = 11.8, p < 0.0001, Small t(14) = 10.3, p<0.0001; L: Large t(3) = 6.78, p = 0.0066). (C) Percentage of mice reaching the L and S criterion in LS group, during the remote test (R) represented as cumulative probability. During the remote test, mice reach the L criterion in just a few days but require more time to reach the S criterion. (D) Performance of individual mice indicating a non-significant reduction in the number of days needed to reach the criterion between the initial and remote training sessions. Only 4 mice (red lines) required more days to reach the criterion during the remote test. (I: 8.067 ± 1.173 %, R: 5.93 ± 0.753, 15 mice; p= 0.204; Mann–Whitney). (E) EdU labeled cell that was also TRAP+ labeled after the initial training but was that was not reactivated (c-Fos) during the remote test. EdU (green), TRAP (red), c-Fos (blue) in TRAP2 mice. Calibration bar, 10 μm. (F) Distribution of TRAP-labeled cells in the IB and SB blades of the dorsal DG in TRAP2 mice. Labeling was performed in LS separation during initial training (I) and compared to control mice that performed only the L configuration. Active neurons show a preferential distribution to the SB over the IB in the DG. This bias is more prominent in the LS group of mice (Two-way ANOVA: group, F(1,12) = 32.98, p < 0.001; blade, F(1,12) = 9.2×10⁻³⁰, p = 1; group×blade interaction, F(1,12) = 16.92, p = 0.0014. Tukey post-hoc test: significant differences were observed between IB; LS and Large (p=0.05), SB; LS and Large (p= 0.05). (G) Distribution of c-Fos labeled cells between the IB and SB blades of the dorsal DG in TRAP2 mice performing the LS separation during remote training (R), compared to control mice that performed only the L configuration (Two-way ANOVA: group, F(1,12) = 56.03, p < 0.00001; blade, F(1,12) = 0, p = 1; group×blade interaction, F(1,12) = 7.02, p = 0.021.Tukey post-hoc test: no significant differences were observed between LS and Large for IB or SB). (H) Percentage of newborn neurons (BrdU (7.7 weeks) or EdU (5.7 weeks)) activated during the initial training (I), as indicated by co-labeling of TRAP+(BrdU+ c-Fos+: L: 0.242 ± 0.242 %, 6 mice, LS: 0.170 ± 0.170 %, 6 mice; ns ; Mann–Whitney). (I) Percentage of newborn neurons (BrdU or EdU) activated during the remote test (R), as indicated by c-Fos expression (BrdU+ c-Fos+: 0.977 ± 0.426 %, 6 mice, p=0.05; EdU+ c-Fos+: 0.694± 0.317 %, 9 mice; p= 0.0785; Mann–Whitney). (J) Percentage of mGCs activated during the (I) training (TRAP+) that were reactivated during the R test (c-Fos+) (expressed as a percentage of TRAP+ cells) (Unpaired t-test, t(6) = 2.90, p = 0.027).

Revised version: (A) Schematic representation of the pattern separation paradigm used in the study. Mice (male n=11, female n=8) performed two trainings sessions: initial (I) and remote (R). The second training (R) took place 3 weeks after the mice reached the S criterion in LS. Mice were perfused when they reached the criterion in S during R. A second group that performed only in the L configuration during both training sessions was included as control (males n=3, females n=3). (B) Learning curves representing the percentage of success during the remote test (R). The control group trained only in the large-separation condition (L group) is shown in green and includes performance in the large-separation trials only. The LS group (orange) was trained in both large-and small-separation conditions and therefore includes two traces: the upper orange trace represents performance during large-separation trials, and the lower orange trace represents performance during small-separation trials. (one-sample t-test to chance at criterion; LS: Large t(14) = 11.8, p < 0.0001, Small t(14) = 10.3, p<0.0001; L: Large t(3) = 6.78, p = 0.0066). (C) Percentage of mice reaching the L and S criterion in LS group, during the remote test (R) represented as cumulative probability. During the remote test, mice reach the L criterion in just a few days but require more time to reach the S criterion. (D) Performance of individual mice indicating a non-significant reduction in the number of days needed to reach the criterion between the initial and remote training sessions. Only 4 mice (red lines) required more days to reach the criterion during the remote test. (I: 8.067 ± 1.173 %, R: 5.93 ± 0.753, 15 mice; p= 0.204; Mann–Whitney). (E) EdU labeled cell that was also TRAP+ labeled after the initial training but was that was not reactivated (c-Fos) during the remote test. EdU (green), TRAP (red), c-Fos (blue) in TRAP2 mice. Calibration bar, 10 μm. (F) Distribution of TRAP-labeled cells in the IB and SB blades of the dorsal DG in TRAP2 mice. Labeling was performed in LS separation during initial training (I) and compared to control mice that performed only the L configuration. Active neurons show a preferential distribution to the SB over the IB in the DG. This bias is more prominent in the LS group of mice (Two-way ANOVA: blade, F(1,12) = 32.98, p < 0.001; group, F(1,12) = 9.2×10⁻³⁰, p = 1; group×blade interaction, F(1,12) = 16.92, p = 0.0014. Tukey post-hoc test: significant differences were observed between IB; LS and Large (p=0.05), SB; LS and Large (p= 0.05). (G) Distribution of c-Fos labeled cells between the IB and SB blades of the dorsal DG in TRAP2 mice performing the LS separation during remote training (R), compared to control mice that performed only the L configuration (Two-way ANOVA: blade, F(1,12) = 56.03, p < 0.00001; group, F(1,12) = 0, p = 1; group×blade interaction, F(1,12) = 7.02, p = 0.021.Tukey post-hoc test: no significant differences were observed between LS and Large for IB or SB). (H) Percentage of newborn neurons (BrdU (7.7 weeks) or EdU (5.7 weeks)) activated during the initial training (I), as indicated by co-labeling of TRAP+(BrdU+ c-Fos+: L: 0.242 ± 0.242 %, 6 mice, LS: 0.170 ± 0.170 %, 6 mice; ns; Mann–Whitney). (I) Percentage of newborn neurons (BrdU or EdU) activated during the remote test (R), as indicated by c-Fos expression (BrdU+ c-Fos+: 0.977 ± 0.426 %, 6 mice, p=0.05; EdU+ c-Fos+: 0.694± 0.317 %, 9 mice; p= 0.0785; Mann–Whitney). (J) Percentage of mGCs activated during the (I) training (TRAP+) that were reactivated during the R test (c-Fos+) (expressed as a percentage of TRAP+ cells) (Unpaired t-test, t(6) = 2.90, p = 0.027).

An analysis of birth-dated active neurons labeled during the initial training (TRAP+) or during the remote trials (c-Fos+) demonstrated a small number of BrdU (∼4 weeks) neurons were labeled during the initial training (Figure 3H) whereas in the remote training we observed both BrdU labeled (7.7 weeks) and EdU labeled (5.7 weeks) abDGCs that were c-Fos+ only in mice after LS training (Figure 3I). No birth-dated abDGCs were found that were double labeled TRAP+ and c-Fos+, however we observed dual labeled mGCs with more TRAP+/c-Fos+ labeled neurons in mice that had performed LS training than those that only performed L training (Unpaired t-test, t(6) = 2.90, p = 0.027) (Figure 3J).

Taken together, these results demonstrate that during initial training and remote training in the pattern separation task, spatially biased activity of neurons is observed in the SB but there is little overlap in activity of neurons in the DG of mice during these epochs of training.

mGC activity is required for DG function and efficient pattern separation in TUNL

Given that the suprapyramidal bias in DG activity scaled with cognitive demand, we hypothesized that altering the excitatory drive within the DG might modulate this spatial pattern of activation. We began by targeting mature granule cells (mGCs), the principal excitatory population in the DG, which are under strong inhibitory control and sparsely active in vivo31,32.

Earlier version: To selectively suppress their activity, Dock10-Cre mice33 were crossed with R26-LSL-Gi-DREADD mice (designer receptor exclusively activated by designer drugs) to express the inhibitory DREADD hM4Di, which can be activated by the ligand deschloroclozapine (DCZ)(50 μg.kg−1)34 during the LS sessions of the task (Figure 4C). Characterization of Dock10-Cre expression in newborn neurons (BrdU+) revealed few labeled cells at 2 weeks postmitosis, with progressively more double-labeled abDGCs in older, non-DCX–positive neurons (Figures 4A and S3), reaching ∼80% of abDGCs by 7 weeks (Figures 4A and 4B). This pattern confirms that Dock10-Cre is expressed primarily in mGCs and older abDGCs. As predicted, DREADD-mediated inhibition of mGCs impaired task performance, with the majority of mice failing to reach criterion (62.5%) within the 14-day LS cutoff, compared to controls (34.8%) (p = 0.0452) (Figures 4D and 4E). Performance in the interleaved low-demand L trials was not affected (not shown). Consistent with reduced DG activity, we observed lower c-Fos expression in both the DG (Unpaired t-test, t(9) = 2.30, p = 0.047) and CA3 (Unpaired t-test, t(9) = 4.93, p = 0.0008) of mice that did not reach criterion after 14 days (Figures 4F, 4G, and 4I). Interestingly, despite the overall reduction in activity, the distribution of active c-Fos+ neurons remained biased toward the SB, similar to control mice that reached criterion (Figures 4H) (two-way ANOVA (unbalanced) (group: F(1,16) = 315.83, p = 5.87 × 10⁻¹² ; blade: F(1,16) = 0, p = 1; group × blade: F(1,16) = 0.040, p = 0.845). This suggests that the overall excitability is a key feature of DG function during performance of a high cognitive demand pattern separation task (Figures 4H, 4J, and 4K).

Revised version: To selectively suppress their activity, Dock10-Cre mice33 were crossed with R26-LSL-Gi-DREADD mice (designer receptor exclusively activated by designer drugs) to express the inhibitory DREADD hM4Di, which can be activated by the ligand deschloroclozapine (DCZ)(50 μg.kg−1)34 during the LS sessions of the task (Figure 4C). Characterization of Dock10-Cre expression in newborn neurons (BrdU+) revealed few labeled cells at 2 weeks postmitosis, with progressively more double-labeled abDGCs in older, non-DCX–positive neurons (Figures 4A and S3), reaching ∼80% of abDGCs by 7 weeks (Figures 4A and 4B). This pattern confirms that Dock10-Cre is expressed primarily in mGCs and older abDGCs. As predicted, DREADD-mediated inhibition of mGCs impaired task performance, with the majority of mice failing to reach criterion (62.5%) within the 14-day LS cutoff, compared to controls (34.8%) (p = 0.0452) (Figures 4D and 4E). Performance in the interleaved low-demand L trials was not affected (not shown). Consistent with reduced DG activity, we observed lower c-Fos expression in both the DG (Unpaired t-test, t(9) = 2.30, p = 0.047) and CA3 (Unpaired t-test, t(9) = 4.93, p = 0.0008) of mice that did not reach criterion after 14 days (Figures 4F, 4G, and 4I). Interestingly, despite the overall reduction in activity, the distribution of active c-Fos+ neurons remained biased toward the SB, similar to control mice that reached criterion (Figures 4H) (two-way ANOVA (unbalanced) (blade: F(1,16) = 315.83, p = 5.87 × 10⁻¹²; group: F(1,16) = 0, p = 1; group × blade: F(1,16) = 0.040, p = 0.845). This suggests that the overall excitability is a key feature of DG function during performance of a high cognitive demand pattern separation task (Figures 4H, 4J, and 4K).

DREADD inhibition of mGCs during a high cognitive demand task.

Earlier version: (A) Percentage of BrdU-labeled cells that are also Dock10+ at different time points following BrdU injections (in days). (B) Representative section showing BrdU+ labeling (green) and Dock10+ labeling (red) in Dock10 cre;Ai9 mice. Calibration bars: 10 μm. (C) Schematic representation of the pattern separation paradigm. All mice (males n = 17, females n = 27) received DCZ i.p. injections 30 minutes before the beginning of their LS trials. (D) Percentage of mice reaching criterion in (S) and percentage that were not successful (NS) after 14 days of training in S configuration of LS. A majority of mice never succeeded in reaching the criterion in the Cre+ DCZ group (z = 2.00, p = 0.0453, z-test - Group NS). (E) Percentage of mice reaching the criterion in S trials across days represented as the cumulative probability including all successful and not successful mice in the group. (F) Density of c-Fos+ GCs in the dorsal DG across the two groups. A significantly reduced density was observed in the Cre+ DCZ group compared to controls (Unpaired t-test, t(9) = 2.30, p = 0.047). (G) Density of c-Fos+ GCs in the dorsal CA3 across the two groups. A significant reduction was observed in the Cre+ DCZ group compared to the control groups (Unpaired t-test, t(9) = 4.93, p = 0.0008). (H) Blade specific distribution of c-Fos+ GCs in the IB and SB of the dorsal DG in mice performing LS separation (Two-way ANOVA (unbalanced, Type II) revealed a significant main effect of group (F(1,16) = 315.83, p = 5.87 × 10⁻¹²), but no significant main effect of blade (F(1,16) = 0, p = 1) and no significant group × blade interaction (F(1,16) = 0.040, p = 0.845). (I) Example of dorsal DG section from Dock10-cre:hM4Di mice after LS training. c-Fos+ neurons are labeled in red. Calibration bars: 100 μm. (J) Spatial distribution of c-Fos+ along the hilar to outer radial GCL axis of the DG (0 to 120μm), with 0 μm indicating location at the SGZ. (K) Cartoon illustrating the distribution of activity labeled GCs in the SB and IB blades of the DG. GCs are depicted in grey while red circles represent activated cells in the DG. The overall distribution of neurons in mice in which mGCs were inhibited (Cre+ DCZ) and which did not reach criterion was not different to controls even though the total density of activity labeled neurons was lower.

Revised version: (A) Percentage of BrdU-labeled cells that are also Dock10+ at different time points following BrdU injections (in days). (B) Representative section showing BrdU+ labeling (green) and Dock10+ labeling (red) in Dock10 cre;Ai9 mice. Calibration bars: 10 μm. (C) Schematic representation of the pattern separation paradigm. All mice (males n = 17, females n = 27) received DCZ i.p. injections 30 minutes before the beginning of their LS trials. (D) Percentage of mice reaching criterion in (S) and percentage that were not successful (NS) after 14 days of training in S configuration of LS. A majority of mice never succeeded in reaching the criterion in the Cre+ DCZ group (z = 2.00, p = 0.0453, z-test - Group NS). (E) Percentage of mice reaching the criterion in S trials across days represented as the cumulative probability including all successful and not successful mice in the group. (F) Density of c-Fos+ GCs in the dorsal DG across the two groups. A significantly reduced density was observed in the Cre+ DCZ group compared to controls (Unpaired t-test, t(9) = 2.30, p = 0.047). (G) Density of c-Fos+ GCs in the dorsal CA3 across the two groups. A significant reduction was observed in the Cre+ DCZ group compared to the control groups (Unpaired t-test, t(9) = 4.93, p = 0.0008). (H) Blade specific distribution of c-Fos+ GCs in the IB and SB of the dorsal DG in mice performing LS separation (Two-way ANOVA (unbalanced, Type II) revealed a significant main effect of blade (F(1,16) = 315.83, p = 5.87 × 10⁻¹²), but no significant main effect of group (F(1,16) = 0, p = 1) and no significant group × blade interaction (F(1,16) = 0.040, p = 0.845). (I) Example of dorsal DG section from Dock10-cre:hM4Di mice after LS training. c-Fos+ neurons are labeled in red. Calibration bars: 100 μm. (J) Spatial distribution of c-Fos+ along the hilar to outer radial GCL axis of the DG (0 to 120μm), with 0 μm indicating location at the SGZ. (K) Cartoon illustrating the distribution of activity labeled GCs in the SB and IB blades of the DG. GCs are depicted in grey while red circles represent activated cells in the DG. The overall distribution of neurons in mice in which mGCs were inhibited (Cre+ DCZ) and which did not reach criterion was not different to controls even though the total density of activity labeled neurons was lower.

DREADD inhibition of abDGCs disrupts blade-biased activity of mGCs and performance of mice in a high cognitive demand task

Earlier version: Having established that mGCs regulate overall DG activity but do not determine the SB-IB bias, we next asked whether manipulating abDGCs could influence this bias and the pattern of DG activation. Because mGC inhibition reduced global excitability without altering the spatial organization of active cells, this suggested that the blade-biased pattern may be set upstream by a more selective cell population. Given that abDGCs constitute a smaller, more plastic excitatory population, we hypothesized that selectively altering their activity might differentially affect both overall excitability and the spatial distribution of active neurons during high cognitive demand tasks. abDGCs are known to affect the excitability of mGCs35,36 and there is a demonstrated blade bias in neurogenesis in the DG. However, it remains unknown whether abDGCs affect the blade-biased activity patterns of mGCs that are evident during a pattern separation task. To acutely modulate the activity of abDGCs during task performance, we created mice in which the inhibitory DREADD (designer receptor exclusively activated by designer drugs) could be conditionally expressed in abDGCs by crossing R26-LSL-Gi-DREADD mice (hM4Di)37 to Ascl1CreERT2 mice. This selectively expresses the hM4Di inhibitory DREADD in a birth-dated cohort of abDGCs by the delivery of TAM (tamoxifen, see Methods) in the diet of mice for 7 weeks allowing the inhibition of their activity with a selective ligand deschloroclozapine (DCZ)(50 μg.kg−1)34 during the LS sessions of the task (Figure 5A & B) 38,39. Inhibition of abDGCs (≤ 7 weeks old) using the DREADD ligand DCZ increased the time animals took to reach criterion compared to control groups of animals that received DMSO (one-way ANOVA, F(2,31) = 3.77, p = 0.034; Tukey post-hoc test indicate that Cre+ DCZ group is significantly different from Cre+ DMSO, while other pairs are not significantly different) (Figure 5C & D). This impairment was specific to the cognitively more demanding LS task, as inhibiting ≤ 7-week-old abDGCs did not affect the performance of mice once they had learned the L configuration (not shown). Importantly control groups that had also received TAM in their diet showed no effect in the pattern separation task or in the biased blade activity of mGCs (Figure 5C-J). Taken together this experiment suggests that abDGCs in this cohort (≤ 7 weeks) are beneficial to performance in a pattern separation task.

Revised version: Having established that mGCs regulate overall DG activity but do not determine the SB-IB bias, we next asked whether manipulating abDGCs could influence this bias and the pattern of DG activation. Because mGC inhibition reduced global excitability without altering the spatial organization of active cells, this suggested that the blade-biased pattern may be set upstream by a more selective cell population. Given that abDGCs constitute a smaller, more plastic excitatory population, we hypothesized that selectively altering their activity might differentially affect both overall excitability and the spatial distribution of active neurons during high cognitive demand tasks. abDGCs are known to affect the excitability of mGCs35,36 and there is a demonstrated blade bias in neurogenesis in the DG. However, it remains unknown whether abDGCs affect the blade-biased activity patterns of mGCs that are evident during a pattern separation task. To acutely modulate the activity of abDGCs during task performance, we created mice in which the inhibitory DREADD (designer receptor exclusively activated by designer drugs) could be conditionally expressed in abDGCs by crossing R26-LSL-Gi-DREADD mice (hM4Di)37 to Ascl1CreERT2 mice. This selectively expresses the hM4Di inhibitory DREADD in a birth-dated cohort of abDGCs by the delivery of TAM (tamoxifen, see Methods) in the diet of mice for 7 weeks allowing the inhibition of their activity with a selective ligand deschloroclozapine (DCZ)(50 μg.kg−1)34 during the LS sessions of the task (Figure 5A & B)38,39. Inhibition of abDGCs (≤ 7 weeks old) using the DREADD ligand DCZ increased the time animals took to reach criterion compared to control groups of animals that received DMSO (one-way ANOVA, F(2,31) = 3.77, p = 0.034; Tukey post-hoc test indicate that Cre+ DCZ group is significantly different from Cre+ DMSO, while other pairs are not significantly different) (Figure 5C & D). This impairment was specific to the cognitively more demanding LS task, as inhibiting

DREADD inhibition of ≤ 7-week abDGCs during a high cognitive demand task.

Earlier version: (A) Schematic representation of the pattern separation paradigm. Prior to onset of the training, mice (males n = 17, females n = 16) were provided tamoxifen-containing food ad libitum in the home cage for 7 weeks to activate TAM inducible Cre recombinase (Ascl1-CreERT2) in abDGCs. All mice received DCZ (deschloroclozapine, 50 μg.kg−1) or DMSO injections 30 minutes before the beginning of LS trials. (B) Example sections from Ascl1-CreERT2; hM4Di mice with immunohistochemical localization of the HA tagged DREADD (top) and after LS training with c-Fos+ neurons labeled in red. Calibration bars: 140 μm. (C) Percentage of mice reaching the criterion in the S trials across days represented as the cumulative probability. The Cre+ DCZ treated group needed additional days to reach criterion compared to controls. (D) Number of days for each mouse to reach the 70% success criterion in S trials, (one-way ANOVA, F(2,31) = 3.77, p = 0.034; Tukey post-hoc test indicate that Cre+ DCZ group is significantly different from Cre+ DMSO, while other pairs are not significantly different). (E) c-Fos+ cell density in all groups of mice. There was an increase in c-Fos+ mGCs in mice in which ≤ 7 week abDGCs were inhibited by DCZ (H(2) = 7.22, p=0.0271; Kruskal Wallis)(Cre+ DCZ comparison to Cre+ DMSO, p = 0.014; Cre+ DCZ comparison to Cre-DCZ, p = 0.027 with Dunn’s multiple comparison post hoc analysis). (F) Example section after c-Fos immunohistochemistry to assess distribution of c-Fos+ mGCs between the IB and SB of the dorsal DG in groups of mice who received DCZ (cre-vs cre+). c-Fos cells are localized closer to the SGZ in the Cre+ group. Calibration bar, 100 μm. (G) c-Fos+ density in the CA3 region of the hippocampus in each of the groups of mice. No significant difference was observed in any of the groups (one-way ANOVA, F(2,13) = 2.27, p = 0.143; Tukey post-hoc test indicates that there is no significant difference between any pairs of groups). (H) Blade distribution of c-Fos+ mGCs in the IB and SB of the dorsal DG in mice performing the LS separation receiving either DCZ or DMSO 30 min prior to task performance. Active neurons are distributed preferentially to the SB than the IB in the DG for the control groups. However this distribution bias is significantly reduced in Cre+ DCZ group (two-way ANOVA (unbalanced) (group: F(1,24) = 659.55, p < 2.2 × 10⁻¹⁶ ; blade: F(2,24) = 3.48 × 10⁻¹³, p = 1; group × blade: F(1,16) = 0.040, p = 0.845). Tukey post-hoc test. Significant differences were observed for IB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765). Significant differences were observed for SB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765). (I) Spatial distribution of c-Fos+ GCs along hilar to molecular layer axes of the SB (0 to 120μm) and the IB (0 to 120μm) of the dorsal DG, with 0 μm indicating the hilar position. c-Fos+ cells are localized closer to the SGZ in the Cre+ DCZ group. (J) Cartoon illustrating the DG with its two blades: IB and SB. Dentate granule cells are depicted in grey while red circles represent activity labeled cells in the DG. For Cre+ DCZ mice, the distribution of labeled neurons is closer to SGZ, away from outer radial GCL and more evenly distributed in IB and SB.

Revised version: (A) Schematic representation of the pattern separation paradigm. Prior to onset of the training, mice (males n = 17, females n = 16) were provided tamoxifen-containing food ad libitum in the home cage for 7 weeks to activate TAM inducible Cre recombinase (Ascl1-CreERT2) in abDGCs. All mice received DCZ (deschloroclozapine, 50 μg.kg−1) or DMSO injections 30 minutes before the beginning of LS trials. (B) Example sections from Ascl1-CreERT2; hM4Di mice with immunohistochemical localization of the HA tagged DREADD (top) and after LS training with c-Fos+ neurons labeled in red. Calibration bars: 140 μm. (C) Percentage of mice reaching the criterion in the S trials across days represented as the cumulative probability. The Cre+ DCZ treated group needed additional days to reach criterion compared to controls. (D) Number of days for each mouse to reach the 70% success criterion in S trials, (one-way ANOVA, F(2,31) = 3.77, p = 0.034; Tukey post-hoc test indicate that Cre+ DCZ group is significantly different from Cre+ DMSO, while other pairs are not significantly different). (E) c-Fos+ cell density in all groups of mice. There was an increase in c-Fos+ mGCs in mice in which:: 7 week abDGCs were inhibited by DCZ (H(2) = 7.22, p=0.0271; Kruskal Wallis)(Cre+DCZ comparison to Cre+ DMSO, p = 0.014; Cre+ DCZ comparison to Cre-DCZ, p = 0.027 with Dunn’s multiple comparison post hoc analysis). (F) Example section after c-Fos immunohistochemistry to assess distribution of c-Fos+ mGCs between the IB and SB of the dorsal DG in groups of mice who received DCZ (cre-vs cre+). c-Fos cells are localized closer to the SGZ in the Cre+ group. Calibration bar, 100 μm. (G) c-Fos+ density in the CA3 region of the hippocampus in each of the groups of mice. No significant difference was observed in any of the groups (one-way ANOVA, F(2,13) = 2.27, p = 0.143; Tukey post-hoc test indicates that there is no significant difference between any pairs of groups). (H) Blade distribution of c-Fos+ mGCs in the IB and SB of the dorsal DG in mice performing the LS separation receiving either DCZ or DMSO 30 min prior to task performance. Active neurons are distributed preferentially to the SB than the IB in the DG for the control groups. However, this distribution bias is significantly reduced in Cre+ DCZ group (two-way ANOVA (unbalanced) (blade: F(1,24) = 659.55, p < 2.2 × 10⁻¹⁶; group: F(2,24) = 3.48 × 10⁻¹³, p = 1; group × blade: F(1,16) = 0.040, p = 0.845). Tukey post-hoc test. Significant differences were observed for IB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765). Significant differences were observed for SB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765). (I) Spatial distribution of c-Fos+ GCs along hilar to molecular layer axes of the SB (0 to 120μm) and the IB (0 to 120μm) of the dorsal DG, with 0 μm indicating the hilar position. c-Fos+ cells are localized closer to the SGZ in the Cre+ DCZ group. (J) Cartoon illustrating the DG with its two blades: IB and SB. Dentate granule cells are depicted in grey while red circles represent activity labeled cells in the DG. For Cre+ DCZ mice, the distribution of labeled neurons is closer to SGZ, away from outer radial GCL and more evenly distributed in IB and SB.

Added in the revised version: ≤ 7-week-old abDGCs did not affect the performance of mice once they had learned the L configuration (not shown). Importantly control groups that had also received TAM in their diet showed no effect in the pattern separation task or in the biased blade activity of mGCs (Figure 5C-J). Taken together this experiment suggests that abDGCs in this cohort (≤ 7 weeks) are beneficial to performance in a pattern separation task.

Prior work has demonstrated that abDGC loss causes an increase in activity of mGCs35,40 which is likely through a loss of connections to hilar interneurons which provide feedback inhibition to the DG. Consistent with this, analysis of c-Fos in the DG demonstrated increased labeling across the DG when abDGCs were inhibited with DCZ activation of hM4Di compared to control groups (p=0.0271) (Figure 5E & F). Conversely, c-Fos+ labeling in the downstream CA3 region of the hippocampus demonstrated no significant difference in density between the control groups and abDGC inhibition (Figure 5G) (one-way ANOVA, F(2,13) = 2.27, p = 0.143; Tukey post-hoc test indicates that there is no significant difference between any pairs of groups).

Earlier version: Interestingly, we found that the biased distribution of c-Fos+ neurons in the two control groups showed the same profile as we had observed before in the IB and SB after mice perform the high cognitive demand task (Figure 5H). However, in the cohort of mice in which abDGCs were inhibited, the prominent biased distribution pattern of activity was not observed (two-way ANOVA (unbalanced) (group: F(1,24) = 659.55, p < 2.2 × 10⁻¹⁶ ; blade: F(2,24) = 3.48 × 10⁻¹³, p = 1; group × blade: F(1,16) = 0.040, p = 0.845). Tukey post-hoc test. Significant differences were observed for IB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765). Significant differences were observed for SB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765) (Figure 5H).

Revised version: Interestingly, we found that the biased distribution of c-Fos+ neurons in the two control groups showed the same profile as we had observed before in the IB and SB after mice perform the high cognitive demand task (Figure 5H). However, in the cohort of mice in which abDGCs were inhibited, the prominent biased distribution pattern of activity was not observed (two-way ANOVA (unbalanced) (blade: F(1,24) = 659.55, p < 2.2 × 10⁻¹⁶; group: F(2,24) = 3.48 × 10⁻¹³, p = 1; group × blade: F(1,16) = 0.040, p = 0.845). Tukey post-hoc test. Significant differences were observed for IB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765). Significant differences were observed for SB comparisons: Cre+ DCZ vs Cre-DCZ (p=0.000548) and Cre+ DCZ vs Cre+ DMSO (p=0.00765) (Figure 5H).

Furthermore, the distribution of c-Fos+ cells within the blades showed that when abDGCs were inhibited during LS training there was a marked prominence of labeled neurons located near the hilus in both the IB and SB in the region closest to the subgranular zone (SGZ), the neurogenic niche of the DG where abDGCs are produced, raising the possibility that these neurons are under greater feedforward inhibitory control via abDGCs (Figure 5I). In summary, DREADD inhibition of ≤ 7-week-old abDGCs during the LS high cognitive demand task resulted in a reduction in both in biased SB-IB spatial activity pattern as well as the activity labeling of neurons along both axes of the transverse DG (Figure 5J), and this altered activity correlated with a reduction in performance of animals which required increased days to complete the high cognitive demand task.

In the previous experiments we had provided TAM to mice in their diet to cause expression of hM4Di in neurons that range from 0 days to 7 weeks. This covers abDGCs that are within their critical period when they demonstrate strong plasticity and excitability and are well connected to the hippocampal network41,42. Because inhibition of this 0–7-week population altered both performance and the blade-biased spatial distribution of mGC activity, we next asked whether these effects depend specifically on the maturity of the abDGCs. To test whether younger, ≤4-week-old neurons contribute similarly to DG function during high cognitive demand, we provided TAM for only 4 weeks to restrict hM4Di expression to these younger abDGCs (Figure 6A). In the LS high cognitive demand task, there was no difference in the time to criterion of animals that had received DCZ administration for hM4Di inhibition of abDGCs compared to the control groups (Figure 6B). Similarly, analysis of c-Fos+ density in the DG found no difference in overall density (Figure 6C)(one-way ANOVA, F(2,13) = 0.157, p = 0.857) and the blade distribution of active mGCs after hM4Di (Figure 6D) or in the hilar to molecular layer axes distribution (Figure 6E).

DREADD inhibition of ≤ 4-week-old abDGCs does not affect performance of mice.

(A) Schematic representation of the pattern separation paradigm used in the study. Prior to the behavioral experiment, mice (males n = 12, females n = 23) were given tamoxifen-enriched food for 4 weeks to activate the Cre recombinase and label newborn neurons. All mice received DCZ or DMSO injections 30 minutes before the beginning of LS trials. (B) Number of days for each mouse to reach the 70% success criterion in S trials (H(2) = 4.21, ns; Kruskal Wallis). (C) Density of c-Fos+ GCs in the dorsal DG across the different groups (one-way ANOVA, F(2,13) = 0.157, p = 0.857). (D) Blade specific distribution of c-Fos+ cells in the IB and SB of the dorsal DG in mice performing LS separation (H(2) = 0.46, ns; Kruskal Wallis). (E) Spatial distribution of c-Fos+ along the hilar to outer radial GCL axis of the DG (0 to 120μm), with 0 μm indicating location at the border of the SGZ and hilus.

Taken together these results demonstrate that ≤ 7-week-old cohort of abDGCs that include neurons in their critical period of development influence the excitability and spatially biased pattern of activity of mGCs in the DG and also affect the performance of mice in a high cognitive demand pattern separation task. In contrast young neurons at ≤ 4 weeks do not influence either the performance of mice in the task or the spatial distribution of activity in the DG.

Discussion

Blade-specific activation patterns in the dorsal dentate gyrus during high-demand discrimination

In this study we used activity labeling techniques combined with a complex, automated visual spatial pattern separation task to capture patterns of activity of GCs in the dorsal DG. The DG is a critical part of the hippocampal formation that receives extrinsic input from cortex and other structures but is also highly regulated by local connections with hilar and CA3 neurons8. While many roles have been ascribed to the dorsal DG, a prominent theory has been that its activity is vital to pattern separation43. We found that activity labeling of neurons demonstrated unique spatial distributions of active GCs which were amplified when the separation task was made more stringent for the mice. We characterized this increased difficulty in the task as an increase in cognitive demand as it reflected the increased difficulty associated with discriminating highly similar spatial locations. However, the change in the task parameters likely engage a complex multitude of behavioral processes, including attention, motivation, reward expectation, and engagement. Therefore, our findings should be interpreted as reflecting activity patterns in th dentate gyrus in response to increased task difficulty rather than a single isolated cognitive operation.

In addition, TRAP and c-Fos labeling integrate neuronal activity across a temporal window, these approaches primarily capture overall task engagement rather than activity linked to specific trial types or behavioral epochs. Accordingly, our findings should be interpreted as reflecting sustained neural recruitment during task performance rather than moment-to-moment encoding of individual trial events. In particular, we observed a biased distribution of active neurons in the SB, which was elevated in mice that completed the LS task. There have been prior descriptions of a blade-biased distribution of active neurons when mice engage in hippocampal dependent behaviors10–13 however it is still not known how these spatially defined activity patterns are modulated or established. We found that not only is there a bias towards more activity in SB and less in the IB during a more demanding task, there is also an increase in activity density in the outer radial areas of the granule cell layer (GCL) where semilunar granule cells are enriched11,30, and in regions closer to the dentate apex. These elevated biased spatial distributions of active neurons were not observed when mice were assessed after running or after the low cognitive demand version of the task, two behaviors that increase the activity of mGCs significantly but do not enhance their blade-biased distribution (Figure S2E). However, while blade-specific differences in activity are evident, no distinct, independent microcircuit has been identified that exclusively governs the function of either the SB or IB. This raises an important question as to whether these differences arise solely from variations in afferent connectivity and intrinsic GC properties or if they reflect an underlying, segregated network organization within the DG. Future studies exploring whether the SB and IB operate as distinct computational units or as interconnected components of a larger network will be essential to clarify the functional basis of this observed bias.

Activity level of mGCs is required for a high demand task

Mature granule cells (mGCs) constitute the primary excitatory population of the DG and are essential for its sparse coding properties. Despite their large number, mGCs are strongly regulated by local inhibition and typically remain sparsely active in vivo31,32. In our study, we used an acute and reversible DREADD inhibition of mGCs using a Cre driver line expressed in mature but not in young adult-born neurons (Figure 4A, S3). Suppressing mGC activity during the LS task had a major effect on performance: most mice in the inhibited cohort failed to reach criterion. Analysis of behavior revealed that during S trials, performance remained at chance, whereas in interleaved L trials, mice continued to perform above criterion. Consistent with reduced mGC excitability, the density of c-Fos+ neurons was significantly decreased in the DG. Interestingly, despite this reduction in overall activation, the blade-biased distribution of active neurons was still evident in mice that performed the task but failed to reach criterion. Thus, unlike the effects observed when inhibiting adult-born neurons, there was no clear correlation between behavioral impairment and changes in the spatial distribution of active cells. Direct inhibition of mGCs therefore produced a fundamentally different outcome on DG circuit dynamics and task performance. While abDGC inhibition slowed performance and altered the spatial distribution of active mGCs even in mice that reached criterion, direct suppression of mGCs reduced global excitability without changing the spatial organization of active cells in animals that did or did not reach criterion (not shown). Moreover, the subset of mice that did reach criterion performed with latencies indistinguishable from control cohorts.

Prior studies using genetic manipulations to increase neurogenesis, have identified blade-specific changes in the number of c-Fos+ cells, with an increase in the IB following contextual discrimination44, and a corresponding increase or decrease in the IB or SB respectively, after navigational learning45. These prior findings support the hypothesis that the SB is context-specific and involved in pattern separation, while the IB is more associated with spatial representation and memory precision46. Our study extends these observations by demonstrating that while overall mGC excitability is required for successful performance in high cognitive demand conditions, the SB–IB bias itself can persist independently of excitatory drive, suggesting that it is not directly determined by mGC activity.

Adult-born granule cells modulate DG network organization during high cognitive demand

The DG, uniquely among temporal lobe structures, contains a neurogenic niche where new neurons are continuously generated and integrated into the existing network. Post-mitotic adult-born dentate granule cells (abDGCs) mature over several weeks47, reaching maturity by around 8 weeks post mitosis where their functional properties and connectivity resemble those of other mGCs48,49. During a critical period between 4–8 weeks, however, they exhibit heightened plasticity and distinct circuit interactions50, making them uniquely positioned to influence local computation.

It has been proposed that abDGCs modulate local circuit activity by contributing to the sparsity of mGC firing through feedback inhibition50–52 or through direct inhibition53. This inhibitory influence is largely mediated by diverse interneuron populations in the DG, including parvalbumin-positive (PV+), somatostatin-positive (SOM+), and neuropeptide Y-positive (NPY+) interneurons, which play crucial roles in regulating the excitability of GCs. Given their importance in enforcing sparse coding, these interneurons could be key modulators of the SB-IB activity differences observed in our study, yet their precise contribution to blade-specific activation remains unclear.

Consistent with prior behavioral work implicating the DG in pattern separation54, manipulations that ablate abDGCs impair discrimination when task difficulty increases55, whereas genetic or environmental manipulations that enhance neurogenesis improve discrimination performance56,57. In our study, only limited numbers of abDGCs were labeled during task performance, as expected given their low population size relative to mGCs. However, neurons that were approximately seven weeks post-mitosis were more prominently labeled than those four weeks post-mitosis, consistent with the notion that older abDGCs are more integrated into the local circuit. Despite their small number, these neurons exerted a disproportionate influence on DG activity.

Using an acute, reversible chemogenetic inhibition of different abDGC cohorts during task performance39, we observed that suppression of ≤7-week-old abDGCs increased mGC activity, consistent with prior studies that chronically reduced neurogenesis35. More importantly, inhibition of abDGCs during the LS task impaired performance in high-demand discrimination trials without affecting already-learned, low-demand trials within the same session. This manipulation disrupted the normal SB-biased pattern of mGC activation seen in controls, as well as the typical distribution of labeled mGCs along the hilar–molecular layer and apex–tip axes of the DG. These effects were specific to the inhibition of ≤7-week-old abDGCs and were not observed when younger (≤4-week-old) or older abDGC populations were targeted. These findings suggest that abDGCs exert their influence not through direct encoding, but through regulating inhibitory tone within the DG microcircuit, potentially via PV-or SOM-positive interneurons that shape granule cell recruitment.

Notably, our results indicate that behavioral performance is not determined solely by the magnitude of dentate gyrus activity. Inhibition of abDGCs increased overall mGC activation while disrupting the spatially organized, blade-biased pattern of activity, whereas direct inhibition of mGCs reduced overall activity without necessarily disrupting spatial organization in animals that remained capable of performing the task. These findings suggest that dentate gyrus function depends on the coordinated recruitment and composition of granule cell ensembles rather than on total activity levels alone. Distinct perturbations may therefore impair behavior by altering the identity, balance, and coordination of neurons participating in task-relevant ensembles.

Together, these results reveal complementary but distinct contributions of mGCs and abDGCs to DG function. mGCs provide the essential excitatory output required for encoding and task performance under high cognitive demand, whereas abDGCs modulate this activity to establish the structured, blade-biased distribution of mGC activation necessary for optimal discrimination. This modulation likely arises through transient inhibitory control during the critical period of abDGC maturation, allowing the DG to dynamically tune its excitability and spatial coding to match task difficulty. Our findings support a model in which efficient pattern separation relies on both the stable excitatory scaffold provided by mGCs and the adaptive, modulatory influence of maturing abDGCs on DG network organization. This aligns with and extends prior models by showing that it is the maturing cohort, not immature or fully mature abDGCs—that provides this modulatory control.

Limitations of the Study

The use of activity tagging has the inherent limitation that it provides a snapshot of activity integrated across a temporal window and therefore cannot resolve which specific trial types, behavioral epochs, or decision processes drive neuronal recruitment. This is countered by the less invasive nature of these experiments than ones that require implants for in vivo recording of neuron activity, which are themselves not fit for purpose because they cannot capture the spatial distribution in the subregions of the DG effectively. Future studies using in vivo measurements of activity will be required to support the present findings and find conclusive neural representations required for discrimination of similar patterns.

Additionally, the experimental design in DREADD experiments required extended duration of tamoxifen administration which resulted in broad labeling of adult-born granule cells spanning multiple stages of maturation (approximately 0–4 weeks and 0–7 weeks old), rather than a precisely birth-dated population. Consequently, the chemogenetic manipulations reflect the combined influence of a heterogeneous cohort of immature neurons, and do not resolve more finely the age-specific contributions of adult-born granule cells. Finally, the circuit mechanisms underlying blade-biased activity remain unknown, and future studies will be required to determine how these spatial patterns arise and how they contribute to computations within the dentate gyrus.

Materials and methods

Ethics statement

All animal procedures were conducted in accordance with protocols approved by the Northwestern University Institutional Animal Care and Use Committee (IACUC). Male and female mice were housed 3–5 per cage on a reverse light/dark cycle of 12/12 hours within a controlled environment. They received food and water ad libitum. All behavioral testing was conducted during the animals’ dark cycle and was conducted by experimenters blinded to genotype. Tail biopsies were collected for genotyping via PCR. All efforts were made to minimize the number of mice used and to reduce any potential pain or stress, in line with ethical guideline.

Animals

All mouse lines were maintained on a C57Bl/6 background. Mice were propagated by female breeders heterozygous for Cre recombinase and homozygous for either the Ai9 allele (Rosa-CAG-LSL-tdTomato-WPRE) or hM4Di allele (R26-LSL-Gi-DREADD) with male breeders homozygous for Ai9 or hM4Di. TRAP2;Ai9 mice were generated by crossing Fos2A-iCreER/+ (TRAP2) mice (Jax #030323)23 with Ai9 (Jax #007909) mice to generate double-heterozygous mice. Ascl1-CreERT2;hM4Di mice were generated by crossing R26-LSL-Gi-DREADD mice (Jax #026219) 37 with mice carrying tamoxifen-inducible Cre recombinase under the control of the Ascl1 promoter (Jax #:012882) 58 to generate double-transgenic progeny Ascl1-CreERT2;R26LSL−hM4Di. Cre-negative littermates from heterozygote breedings were used as controls. Dock10-Cre;hM4Di mice were generated by crossing Dock10-Cre 33 mice with hM4Di mice to produce double-transgenic progeny Dock10-Cre;R26LSL−hM4Di.

Birth-dating of abDGCs

6–8-week-old mice received three intraperitoneal (i.p.) injections of 5-Bromo-2′-deoxyuridine (BrdU, Sigma Aldrich #B9285) at 100 mg/kg (i.p) at 4 hours intervals on a single day 59. Two weeks later, 5-Ethynyl-2′-deoxyuridine (EdU, Sigma Aldrich, #900584) was administrated following the same injection protocol. EdU was used as an additional thymidine analog to label a distinct cohort of dividing cells. Because BrdU and EdU are incorporated into DNA during S-phase following systemic administration, the birth-dating timeline and survival intervals applied to BrdU labeling also apply to EdU labeling.

In this study, mature granule cells (mGCs) are operationally defined as dentate granule cells that were not labeled with BrdU or EdU and therefore are not classified as adult-born neurons within the defined labeling window. While the majority of these cells are expected to represent developmentally generated mature neurons, we cannot exclude the possibility that a small fraction may include younger, unlabeled neurons generated outside the labeling period.

Tamoxifen administration

IP injections: To time the labeling of TRAP+ cell in TRAP2 mice, 4-hydroxytamoxifen (4-OHT; Sigma Aldrich, #H7904) was administered intraperitoneally. 4-OHT was dissolved at 20 mg/mL in pure ethanol and stored at –20°C for up to several weeks. Before use, 4-OHT was redissolved in corn oil (Sigma, Cat #C8267) and heated to 90°C with shaking until the ethanol had evaporated yielding a final concentration of 10 mg/mL. The final 10 mg/mL 4-OHT solutions was used on the day it was prepared. All injections were delivered intraperitoneally at 50 mg/kg either at the end of the training phase of the day or 2 hours after running.

Added in the revised version: Oral administration

Earlier version: Oral administration: For conditional expression of hM4Di in Ascl1-CreERT2 cells, a tamoxifen diet (TD.130858, Envigo) was provided ad libitum for 4 or 7 weeks, depending on the experimental timeline. Body weight was carefully monitored daily. The chow contained 500 mg tamoxifen/kg, delivering ∼80 mg tamoxifen/kg body weight per day, assuming an average mouse weight of 20-25 g and a daily intake of 3-4 g.

Revised version: For conditional expression of hM4Di in Ascl1-CreERT2cells, a tamoxifen diet (TD.130858, Envigo) was provided ad libitum for 4 or 7 weeks, depending on the experimental timeline. Body weight was carefully monitored daily. The chow contained 500 mg tamoxifen/kg, delivering ∼80 mg tamoxifen/kg body weight per day, assuming an average mouse weight of 20-25 g and a daily intake of 3-4 g.

Mice were monitored daily for signs of stress, discomfort, or adverse effects related to tamoxifen exposure.

Inhibitory DREADD-hM4Di activation

Deschloroclozapine (DCZ, MedChemExpress #HY-42110) was dissolved in 2% DMSO in 0.9% saline and administered by i.p. injections at a dose of 50 μg.kg−1 34. Behavioral trials started 30 minutes after drug administration.

Behavioral procedures

Adult (8-10 weeks-old) male and female mice were used in behavioral studies. To acclimate the mice to handling before behavioral testing, all animals were handled by the experimenter for 2 min per day for 3 consecutive days.

Running wheels

To habituate the mice to running, wireless running wheels (ENV-047, Med Associates) were placed in their home cage for 5 days (maximum 5 mice per cage). Two days prior to BrdU injections, the wheels were removed from the cage. One week before the experiment, mice were individually housed. On the day of the experiment, a running wheel was added to each cage for a duration of 4 hours. After 2 hours, mice received a 4-OHT injection and were returned to their respective cages. Following an additional 2 hours of running, the wheels were removed from the cages. Three days later, wheels were reintroduced to the cage and remained until mice were euthanized for assessment of TRAP+ cells, 90 minutes after the beginning of the run. The effect of exposure to the wheels was quantified by measuring the total distance run by each mouse.

Automated touch screen operant behaviors

Testing was conducted in custom-built, touch screen-based automated operant device boxes designed specifically for mice, the Operant House (doi.org/10.1101/2025.01.24.634815). The apparatus consisted of a rectangular modular testing chamber (w14.5cm x l21.5cm x h18.5cm). A 60mW green LED was located above a port that gave access to a retractable water bottle at the back of the chamber. At the opposite end of the chamber, a flat-screen infrared touch monitor was controlled by custom scripts (Python). In front of the touchscreen was a mask with 2 horizontal lines of 5 squares. Mice performed a modified version of the “trial unique nonmatching-to-location” (TUNL) task 29. Five days prior to testing and during the training phase, mice were placed on a water-restriction regime, with access to water for 10 minutes per day, and were maintained at 85-90% body weight during the entire duration of the experiment. Mice were trained each day, in two separate sessions of 30 trials or a maximum of 30 minutes.

Training phases – Shaping: Following 30 minutes of habituation to the operant chamber, mice were trained to touch a lit white square stimulus presented randomly in one of the response windows of the touch screen in order to receive a reward (2 seconds water nozzle access with saccharin water 0.01%). Mice were trained in Shaping until they successfully achieved 30 rewards within 30 min during 3 successive sessions.

Training phases – Large-only separation (low cognitive demand): After shaping, mice were trained in the TUNL task with a Large-only separation (L configuration). For this task, each trial was composed of two phases. During the sample phase, a white square was illuminated in one of the 10 possible locations on the screen (Figure 1A, S2A & B). Following a nose-poke to the sample square the stimulus disappeared, and the mouse received a reward only on 33% of the trials. After a 1s delay from the sample trial, the choice trial was initiated where two stimuli were presented one in the repeated sample location (incorrect) and the other in the new location (correct) separated spatially by 4 squares in the L configuration (Figure 1A, S2A & B). A touch to the correct sample location resulted in a reward and an 5s inter-trial interval (ITI) before the next trial was initiated. An incorrect choice resulted in a 10 second punishment (roof lights ON) and then an ITI, followed by correction trials during the first two days of training. Correction trials followed the same procedures, except that the same sample and choice locations from the previous incorrect trial were presented again until the correct choice was made. The correction trials were not included in performance calculations. Mice progressed to the next training phase after completing ≥30 trials within 30 minutes and achieving >70% correct responses in at least two of four sessions across two consecutive days.

Training phases – Large + Small separation (high cognitive demand): Mice were then trained in sessions containing both large-separation (L) and small-separation (S) trials (LS configuration). Training followed the same pattern as previously described except that there were two separate possibilities for the presentation of the choice square which was either separated by 4 squares from the sample square during L trials, or the choice square was 0-1 square separated from the sample location during S trials (Figure 1A & S2B).

Each session consisted of 15 L and 15 S trials presented in alternating order. Mice were trained until reaching criterion defined as ≥70% correct responses on S trials in two of four sessions across two consecutive days.

Approximately 25–30% of mice across experimental groups did not reach LS criterion within the allotted training period (Figure S2G).

Immunohistochemistry

Mice were deeply anesthetized with isoflurane and then perfused transcardiacly with PBS containing 0.02% sodium nitrite followed by 2% paraformaldehyde (PFA) in 0.1 M sodium acetate buffer (pH 6.5) for 3 min, and 2% PFA in 0.1 M sodium borate buffer (pH 8.5) for 10 min 60. Brains were dissected and post-fixed in 2% PFA 0.1 M sodium borate buffer (pH 8.5), for 24-48 hours and stored at 4°C in PBS until sectioning. 50μm floating sections were collected on a Leica Vibratome VT1000 and stored at 4°C in PBS containing 0.02% sodium azide. For c-Fos analysis, perfusions were performed 90 minutes after the end of the behavioral training. Immunohistochemistry was conducted at room temperature using standard procedures. Free-floating sections were incubated for 30 min in a solution of 25 mM glycine to decrease the background signal and then rinsed two times in Tris-Buffered Saline 1X (TBS). Sections were then blocked for 1 hour in the saturation solution containing TBS-Triton (0.25% Triton X-100 (Fisher, BP151-100), 3% normal donkey serum (NDS; Jackson ImmunoResearch, 017-000-121) and 0.1% bovine serum albumin (BSA; Sigma, A9085)). The sections were then incubated overnight in the saturation solution containing the primary antibody (rabbit anti c-Fos (1/1500, Cell Signaling), sheep anti-BrdU (1/2500, NovusBio), rabbit anti RFP (1/5000, Abcam), rabbit anti HA-Tag (1:500, #3724, Cell Signaling). For c-Fos staining, no serum and BSA were used. For BrdU staining, sections were incubated for 10 mins in a solution of DNAse I (Sigma #D5025, 15 KU/mL dissolved in 0.05 M NaCl), MgCl2 (6 mM) and CaCl2 (10 mM) to cleave DNA then rinsed 2 x 5 min in TBS before the saturation part. For EdU staining, sections were incubated in solution from the kit ClickITTM (#C10340 Thermofisher) during 30 min then rinsed 2 x 5 min TBS before the saturation part. No antibodies were necessary for the rest of the immunostaining. The next day, sections were washed 2×5 min in TBS before being incubated for 1 hour in secondary antibody solution (donkey anti-rabbit Alexa 647 or 488 (1/1000, Invitrogen), donkey anti-sheep Alexa 488 (1/1000, Invitrogen)). Finally, sections were washed 2×5 min in TBS and cover-slipped with ProLong Diamond mounting media 275 containing DAPI (#P36962, ThermoFisher Scientific).

Imaging and neuron quantification

Analyses were conducted in the dorsal DG, as prior studies have demonstrated its involvement in spatial memory and the fine-grained contextual differentiation required for pattern separation. The dorsal hippocampus is crucial for encoding and distinguishing spatial information, which is central to this process. In contrast, the ventral DG is more strongly associated with emotional processing and affective memory rather than high-resolution spatial encoding needed for pattern separation.

Confocal images were obtained with a Zeiss microscope (LSM700) with a 20x objective with the experimenter blind to the genotype and condition following previous protocols 59. One section in every ten serial sections were analyzed through the dorsal hippocampus (stereotaxic coordinates:-1.20 to-2.30 relative to bregma according to Paxinos’ brain atlas). Two to three 10 μm z-stack acquisitions, were made to picture the whole dentate gyrus in each section. 3D reconstructions were used to univocally verify co-localization (e.g. activated mature neurons or immature neurons). All acquisitions were carried out in sequential scanning mode to prevent cross-bleeding between channels. Stacks of images were then reconstituted using Fiji software, allowing the experimenter to count the number of cells. The surface area of the dorsal hippocampus, DG, SB, IB and CA3 were traced in mapping software (Neurolucida Zeiss) using a 10x objective. The volume was determined by multiplying the surface area by the distance between sections (500 μm). Densities of BrdU+, EdU+, TRAP+ and c-Fos+ cells were estimated by multiplying the total number of labeled cells by 10. For immunohistochemical analysis, a subset of approximately 5 mice from the behavioral cohorts of ∼15 mice were randomly selected. If the results exhibited high variability, one or two additional mice were included to improve data precision. Selection remained random throughout. This approach was chosen to balance the workload associated with immunohistochemistry while ensuring sufficient statistical power; given that behavioral experiments require larger sample sizes to achieve robust significance only a subset could be included in this analysis.

Distribution of TRAP+ and c-Fos+ Cells in the DG

We followed methods presented in literature 61. Using stacks of images obtained from confocal microscopy, the localization of TRAP+ and c-Fos+ cells in the DG was assessed using Fiji software. Distance from the subgranular zone (SGZ): The contours of each blade were manually outlined and the distance from the center of each cell body to the SGZ was manually measured. Cells located in the apex region were excluded from the counting.

Distance from the apex: Excluding the apex region, the size of each suprapyramidal blade was measured as well as the distance of each cell from the apex. These distances were then expressed as percentage of total blade length.

Statistical analysis

Data are expressed as mean ± SEM. Statistical analyses were performed using Statview and Origin. Histochemical and other experimental data were assessed for normality using the Shapiro–Wilk test. For normally distributed datasets, comparisons between two groups were performed using unpaired t-tests, while one-way or two-way ANOVAs were used for experiments with multiple factors, followed by appropriate post hoc tests (Tukey or Bonferroni) when significant. For datasets that deviated from normality, non-parametric tests (Mann–Whitney U test or Kruskal–Wallis test) were applied, with Dunn’s post hoc test and Bonferroni correction as appropriate. Differences were considered significant at p ≤ 0.05.